<question type="stack">
    <name>
      <text>Quotienten- und Kettenregel Ableitung</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Bestimmen Sie die erste Ableitung der folgenden Funktion mit Hilfe der Ketten- und Quotientenregel:&lt;br&gt;&lt;/p&gt;&lt;div class="editor-indent" style="margin-left: 30px;"&gt;&lt;p style=""&gt;\( f(x)= \dfrac{e^{{@A@}x^2+1}}{{@B@}x^2} \)&lt;/p&gt;\( f'(x) = \)&amp;nbsp;&amp;nbsp; [[input:ans1]] [[validation:ans1]][[feedback:prt1]]&lt;/div&gt;&lt;p&gt;&lt;/p&gt;&lt;p&gt;&lt;/p&gt;&lt;p&gt;&lt;/p&gt;&lt;p&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;p&gt;&lt;/p&gt;&lt;h4&gt;&lt;strong&gt;Lösung&lt;/strong&gt;&lt;/h4&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;\(f(x)=\dfrac{e^{{@A@}x^2+1}}{{@B@}x^2}\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;Sei \(u(x)=e^{{@A@}x^2+1}\)&amp;nbsp; &amp;nbsp; mit&amp;nbsp; &amp;nbsp; \(u^{\prime}(x)={@2*A@}xe^{{@A@}x^2+1}\) (Kettenregel)&lt;/p&gt;&lt;p&gt;und \(v(x)={@B@}x^2\)&amp;nbsp; &amp;nbsp; mit&amp;nbsp; &amp;nbsp; \(v^{\prime}(x)={@2*B@}x\)&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;Mit&amp;nbsp;&lt;/p&gt;&lt;p&gt;\(f(x)=\dfrac{u(x)}{v(x)}=\dfrac{e^{{@A@}x^2+1}}{{@B@}x^2}\)&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;und der Quotientenregel folgt&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;\(f^{\prime}(x)=\dfrac{u^{\prime}(x)v(x)-u(x)v^{\prime}(x)}{v(x)^2}\)&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp;\(=\dfrac{{@2*A@}xe^{{@A@}x^2+1}\cdot{@B@}x^2-e^{{@A@}x^2+1}\cdot{@2*B@}x}{({@B@}x^2)^2}\)&amp;nbsp; &amp;nbsp;&amp;nbsp;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp;\(=\dfrac{({@2*A*B@}x^3-{@2*B@}x)e^{{@A@}x^2+1}}{{@B^2@}x^4}\)&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp; &amp;nbsp;\(=\dfrac{({@2*A*B@}x^2-{@2*B@})e^{{@A@}x^2+1}}{{@B^2@}x^3}\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp;&amp;nbsp;&lt;/p&gt;&lt;p&gt;&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>A:2+rand(9);
B:2+rand(9);
f:e^(A*x^2+1)/(B*x^2);

m1:diff(f,x);
m2:ratsimp(m1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text />
    </specificfeedback>
    <questionnote>
      <text>n</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>m2</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>getestet</text>
</tag>
      <tag><text>Lösung</text>
</tag>
      <tag><text>Quotientenregel</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
      <tag><text>Übung</text>
</tag>
      <tag><text>Kettenregel</text>
</tag>
    </tags>
  </question>

