<question type="stack">
    <name>
      <text>Ableitung von Polynomen</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Wie lautet die Ableitung der folgenden Funktion?&lt;/p&gt;
&lt;p&gt;\[f(x)={@myf@}\]&lt;/p&gt;
&lt;p&gt;Validieren Sie Ihre Eingabe vor Abgabe durch Klick auf den Haken.&lt;/p&gt;
&lt;p&gt;\(f'(x)=\) [[input:ans1]] [[validation:ans1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text />
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>/* Zufall */
a1:rand_with_step(2,9,1);
a2:rand_with_step(2,9,1);
a3:rand_with_step(2,9,1);
b1:rand_with_step(7,9,1);
b2:rand_with_step(4,6,1);b3:rand_with_step(2,3,1);
/* Funktion */
myf:ev(a1*x^(b1)+a2*x^(b2)+a3*x^(b3),simp);
f(x):=myf;
/* Antwort */
ta1:diff(f(x),x,1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </specificfeedback>
    <questionnote>
      <text>{@myf@}</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text />
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text />
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text />
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>getestet</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
    </tags>
  </question>

