<question type="stack">
    <name>
      <text>011) Differenzenquotient</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Wir betrachten die Funktion \(f\), die durch folgende Gleichung definiert ist:&lt;/p&gt;
&lt;p&gt;\[f(x)={@tshow@}\]&lt;/p&gt;
&lt;p&gt;Wie lautet der Differenzenquotient \[\frac{f(x+h)-f(x)}{h}\] an der Stelle \(x={@x1@}\)?&lt;/p&gt;
&lt;p&gt;[[input:ans1]] [[validation:ans1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text />
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>/* Zufallszahlen */
a1:rand_with_step(2,5,1);
b1:rand_with_step(1,9,1);
a2:rand_with_step(6,9,1);
b2:rand_with_step(1,9,1);
a3:a1;
b3:rand_with_step(1,9,2);
a4:a2;
b4:rand_with_step(1,9,2);
x1:rand_with_step(1,9,1);
/* mögliche Antworten */
f1:a1*x^2+b1; 
f2:a2*x^2+b2; 
f3:a3*x^2+b3*x; 
f4:a4*x^2+b4*x; s1:ev(a1*h+2*a1*x1,simp);
s2:ev(a2*h+2*a2*x1,simp);
s3:ev(a3*h+2*a3*x1+b3,simp);
s4:ev(a4*h+2*a4*x1+b4,simp);
/* wähle richtige Antwort zufällig aus */
myterms:[[f1,s1], [f2,s2], [f3,s3], [f4,s4]];fshuffle:rand_selection(myterms, 4);
tshow:fshuffle[1][1];
tac:fshuffle[1][2];
/* alle Antwortmöglichkeiten */
ta: [[fshuffle[1][2], true], [fshuffle[2][2], false], [fshuffle[3][2], false], [fshuffle[4][2], false]];
ta:random_permutation(ta);
tao:[null, false, "Der korrekte Ausdruck ist hier nicht aufgelistet."];
ta:append(ta,[tao]);
/* Technik */
tafunc:maplist(first, ta);
tabool:maplist(second, ta);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </specificfeedback>
    <questionnote>
      <text>{@tshow@}</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text />
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text />
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text />
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>setans:setify(ans1);</text>
      </feedbackvariables>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>elementp(tafunc[1], setans)</sans>
        <tans>tabool[1]</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <answertest>AlgEquiv</answertest>
        <sans>elementp(tafunc[2], setans)</sans>
        <tans>tabool[2]</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <answertest>AlgEquiv</answertest>
        <sans>elementp(tafunc[3], setans)</sans>
        <tans>tabool[3]</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <answertest>AlgEquiv</answertest>
        <sans>elementp(tafunc[4], setans)</sans>
        <tans>tabool[4]</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <answertest>AlgEquiv</answertest>
        <sans>elementp(tafunc[5], setans)</sans>
        <tans>tabool[5]</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>getestet</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
      <tag><text>Differenzenquotient</text>
</tag>
    </tags>
  </question>

