<question type="stack">
    <name>
      <text>IM1 (16) - Tangente parallel zu Gerade</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Gegeben ist die Funktion &lt;span class="latex"&gt;f: \mathbb{R} \to \mathbb{R}&lt;/span&gt; mit \(f(x)=e^{{@a@}x}\). Bestimmen Sie die Koordinate &lt;span class="latex"&gt;x_0&lt;/span&gt; des Punktes, in dem die Tangente von &lt;span class="latex"&gt;f&lt;/span&gt; parallel zur Geraden \(y={@m@}x{@c@}{@b@}\) verläuft.&lt;/p&gt;
&lt;p style="text-align: right;"&gt;(3 Punkte)&lt;/p&gt;
&lt;hr /&gt;
&lt;p&gt;&lt;em&gt;Geben Sie das Ergebnis als Dezimalzahl in dem entsprechenden digitalen Antwortfeld ein. Runden Sie gegebenenfalls &lt;strong&gt;auf genau 4 Nachkommastellen&lt;/strong&gt;.&lt;br /&gt;&lt;em&gt;Benutzen Sie &lt;strong&gt;Punkt&lt;/strong&gt; statt Komma, um die Nachkommastellen zu trennen.&lt;br /&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;Mit dem Haken neben dem Eingabefeld können Sie überprüfen, wie Ihr Ergebnis vom System gelesen wird.&lt;/em&gt;&lt;br /&gt;Beispiele für Eingaben: &lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/em&gt;\(3x \rightarrow 3x, \log(3x) \rightarrow \ln(3x), x \hat{} 3 \rightarrow x^3, x \hat{} (1/3) \rightarrow \sqrt[3]{x}, sqrt(3) \rightarrow \sqrt{3}, pi \rightarrow \pi \)&lt;/p&gt;
&lt;p&gt;&lt;br /&gt;&lt;em&gt;&lt;/em&gt;&lt;/p&gt;
&lt;p style="text-align: left;"&gt;&lt;em&gt; &lt;/em&gt;&lt;strong&gt;Antwort:&lt;/strong&gt;&lt;/p&gt;
&lt;p&gt;\(x_0=\)[[input:ans1]] [[validation:ans1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text />
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text />
    </stackversion>
    <questionvariables>
      <text>/* e^(a*x) */a: rand([3,4,6,8]);/* y=mx+b */m: rand([12,16,24,36]);b: rand([1,2,3]);c: if b=1 then "+" else "-"; /* c=Vorzeichen von b, da +1,-2,-3 möglich sein sollen (nur relevant für Aufgabenstellung) *//* erste Ableitung */f1: a*e^(a*x);/* f1=b, Gleichung nach x auflösen */x: (1/a)*log(m/a);/* float macht aus derLösung eine Gleitkommazahl, fpprintprec gibt die Anzahl Nachkommastellen vor*/fpprintprec : 4;M_ans1: float(x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </specificfeedback>
    <questionnote>
      <text>\(f(x)=e^{{@a@}x}\) und \(y={@b@}x{@d@}{@c@}\)</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct answer, well done.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;Your answer is partially correct.&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;Incorrect answer.&lt;/p&gt;</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>M_ans1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>M_ans1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>ungetestet</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
      <tag><text>Tangente</text>
</tag>
    </tags>
  </question>

