<question type="stack">
    <name>
      <text>022) Monotonie, stationäre Stellen</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;&lt;/p&gt;&lt;p&gt;Sei die folgende Funktion gegeben:&lt;br&gt;&lt;/p&gt;&lt;div class="editor-indent" style="margin-left: 30px;"&gt;&lt;p&gt;\(f(x)={@f@}\)&lt;/p&gt;&lt;p&gt;{@plot([f,fa,faa], [x,2*b1,-2.5*b1], [y,(c*(-b1)^a1+c1*b*(-b1)^a)-20,-(c*(-b1)^a1+c1*b*(-b1)^a)+20],[box, false], [yx_ratio, 0.7], [ytics,(c*(-b1)^a1+c1*b*(-b1)^a),(-(c*(-b1)^a1+c1*b*(-b1)^a))/2,-(c*(-b1)^a1+c1*b*(-b1)^a)], [axes, solid],[legend,"f(x)","f'(x)","f''(x)"])@}&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;/p&gt;&lt;/div&gt;&lt;p&gt;&lt;b&gt;(a)&lt;/b&gt; \(f^{\prime}(x)=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;b&gt;Hinweis: &lt;/b&gt;Geben Sie gegebenenfalls&lt;b&gt; inf &lt;/b&gt;für \(\infty\) ein.&lt;/p&gt;&lt;p&gt;&lt;b&gt;&lt;b&gt;(b)&lt;/b&gt;&amp;nbsp;&lt;/b&gt;Geben Sie die maximalen Intervalle \(I_1\)&amp;nbsp; und \(I_2\) an, in denen \(f\)&lt;b&gt; streng monoton fallend&amp;nbsp;&lt;/b&gt;ist&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; \(I_1=\Big]\) [[input:ans2]] \(,\) [[input:ans3]] \(\Big[\) [[validation:ans2]][[validation:ans3]]&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; \(I_2=\Big]\) [[input:ans4]] \(,\) [[input:ans5]] \(\Big[\) [[validation:ans4]][[validation:ans5]][[feedback:prt2]]&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;b&gt;&lt;/b&gt;&lt;b&gt;(c)&lt;/b&gt;&amp;nbsp;Geben Sie das maximale Intervall \(I\) an, in dem \(f\)&amp;nbsp;&lt;b&gt;streng monoton steigend&lt;/b&gt;&amp;nbsp;ist&lt;br&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; \(I=\Big]\) [[input:ans6]] \(,\) [[input:ans7]] \(\Big[\) [[validation:ans6]][[validation:ans7]][[feedback:prt3]]&lt;/p&gt;&lt;br&gt;&lt;p&gt;&lt;/p&gt;&lt;p&gt;&lt;b&gt;(d)&lt;/b&gt;&amp;nbsp;\(f^{\prime \prime}(x)=\) [[input:ans8]] [[validation:ans8]][[feedback:prt4]]&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;b&gt;(e)&lt;/b&gt; Bei \(x=\) [[input:ans9]] hat \(f(x)\) ein/einen [[input:ans10]].&amp;nbsp;[[validation:ans9]][[validation:ans10]]&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp;Bei \(x=\) [[input:ans11]] hat \(f(x)\) ein/einen [[input:ans12]].&amp;nbsp;[[validation:ans11]][[validation:ans12]][[feedback:prt5]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h4&gt;&lt;strong&gt;Lösung&lt;/strong&gt;&lt;/h4&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;strong&gt;(a)&lt;/strong&gt;&amp;nbsp;\(f^{\prime}(x)={@fa@}={@c1@}x^{@a-1@}(x{@b1@})\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;strong&gt;&lt;strong&gt;(b) &lt;/strong&gt;&lt;/strong&gt;\(f\) ist&lt;strong&gt;&amp;nbsp;streng monoton fallend&lt;/strong&gt;, wenn \(f^{\prime}(x)&amp;gt;0\) gilt.&lt;br&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; Das ist erfülllt, wenn \(x&amp;lt;{@-b1@}\) und \(x \neq 0\):&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; \(I_1=]-\infty,0[\)&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; \(I_2=]0,{@-b1@}[\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;strong&gt;(c)&lt;/strong&gt;&amp;nbsp; \(f\) ist&amp;nbsp;&lt;strong&gt;streng monoton steigend&lt;/strong&gt;, wenn \(f^{\prime}(x)&amp;lt;0\) gilt.&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; Das ist erfülllt, wenn \(x&amp;gt;{@-b1@}\):&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; \(I=]{@-b1@}, \infty[\)&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;strong&gt;(d)&lt;/strong&gt;&amp;nbsp;\(f^{\prime \prime}(x)={@faa@}\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;/p&gt;&lt;p&gt;&lt;strong&gt;(e)&lt;/strong&gt;&amp;nbsp;Bei \(x=0\) hat \(f(x)\) einen Sattelpunkt.&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; (\(f^{\prime}(0)=0=f^{\prime \prime}(0)\))&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp;Bei \(x={@-b1@}\) hat \(f(x)\) ein Minimum.&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp; (\(f^{\prime}({@-b1@})=0\) und \(f^{\prime \prime}({@-b1@})={@fb@}&amp;gt;0\))&lt;br&gt;&lt;/p&gt;&lt;br&gt;&lt;p&gt;&lt;/p&gt;&lt;br&gt;&lt;p&gt;&lt;/p&gt;&lt;p&gt;&amp;nbsp; &amp;nbsp; &amp;nbsp;&amp;nbsp;&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>a: rand([3,7])
a1: a+1
b: rand([-4-3,-2])
c: rand_with_prohib(2,5,[a,a1,-b])
b1: b*a
c1: c*a1

fa: c1*x^a+c1*b1*x^(a-1)
f: int(fa,x)
faa: diff(fa,x)

Tans1 : 1
AnsList1: [[1, true, "Minimum"], [2, false, "Maximum"],[3, false, "Sattelpunkt"]];
Tans2 : 3
AnsList2: [[1, false, "Minimum"], [2, false, "Maximum"],[3, true, "Sattelpunkt"]];

fb: ev(faa, x=-b1)</text>
    </questionvariables>
    <specificfeedback format="html">
      <text />
    </specificfeedback>
    <questionnote>
      <text>Monotonie</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff(f,x)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans10</name>
      <type>dropdown</type>
      <tans>AnsList1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans11</name>
      <type>algebraic</type>
      <tans>0</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans12</name>
      <type>dropdown</type>
      <tans>AnsList2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>-inf</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>0</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>0</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>-b1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>-b1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>inf</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>diff(diff(f,x),x)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>-b1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff(f,x)</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>-inf</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>5</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>0</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>-b1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>3</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>-inf</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>4</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>0</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>0</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>6</truenextnode>
        <trueanswernote>prt2-6-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-6-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>0</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>7</truenextnode>
        <trueanswernote>prt2-7-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-7-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>-b1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-8-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-8-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>-b1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>inf</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>diff(diff(f,x),x)</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>-b1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>Tans1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>2</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>0</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>3</truenextnode>
        <trueanswernote>prt5-3-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-3-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>Tans2</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-4-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-4-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>0</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>5</truenextnode>
        <trueanswernote>prt5-5-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-5-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>Tans2</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty />
        <truenextnode>6</truenextnode>
        <trueanswernote>prt5-6-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-6-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>-b1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty />
        <truenextnode>7</truenextnode>
        <trueanswernote>prt5-7-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-7-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>Tans1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-8-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-8-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>getestet</text>
</tag>
      <tag><text>Lösung</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
      <tag><text>Extrema</text>
</tag>
      <tag><text>Monotonie</text>
</tag>
      <tag><text>Plot</text>
</tag>
      <tag><text>Transfer</text>
</tag>
    </tags>
  </question>

