<question type="stack">
    <name>
      <text>Newton-Verfahren</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Um den Effektivzins eines Ratenkredits zu ermitteln, kommt das Newton-Verfahren zur Anwendung.&lt;/p&gt;
&lt;p&gt;Wir betrachten hierfür einen Ratenkredit über {@kredit@} Euro mit einer Laufzeit von {@lfzmonate@} Monaten bei einem Monatszins von {@iPrint@}%. Es ist eine monatliche (nachschüssige) Rate in Höhe von {@rate@} zu zahlen, so dass der Kredit am Ende der Laufzeit getilgt ist. Wenn der Kreditgeber zusätzlich eine einmalige Bearbeitungsgebühr bei Kreditaufnahme in Höhe von {@gebuehr@} Euro berechnet, dann ist der effektive Monatszins eine Nullstelle der folgenden Gleichung:&lt;/p&gt;
&lt;p&gt;\[{@eqn@}\]&lt;/p&gt;
&lt;p&gt;Wir nutzen die linke Seite der Gleichung als Definition einer Funktion \(f\) mit:&lt;/p&gt;
&lt;p&gt;\[f(x)={@lhs(eqn)@}\]&lt;/p&gt;
&lt;p&gt;Führen Sie den ersten Schritt des Newton-Verfahrens für diese Funktion durch. Verwenden Sie als Startwert den Wert \(x_1={@xstart@}\) mit \(f({@xstart@})={@fxstart@}\).&lt;/p&gt;
&lt;p&gt;Runden Sie das Ergebnis jeweils kaufmännisch auf fünf Stellen. Verwenden Sie als Dezimaltrennzeichen bei allen Aufgaben einen Punkt.&lt;/p&gt;
&lt;p&gt;a) Bestimmen Sie nun den ersten Wert des Newton-Verfahren für diese Funktion.&lt;/p&gt;
&lt;p&gt;\(x_2=\)[[input:ans1]] [[validation:ans1]]&lt;/p&gt;
&lt;p&gt;b) Wie groß ist der Fehler an dieser Stelle?&lt;/p&gt;
&lt;p&gt;\(\left|f(x_2)\right|=\)[[input:ans2]] [[validation:ans2]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text />
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>/* Zufall */
kredit:rand_with_step(10000,20000,1000);
gebuehr:rand_with_step(500,1500,100);
monatszins:rand_with_step(.005,.015,.001);
lfzjahre:rand_with_step(3,6,1);
/* Berechnungen */
iPrint:monatszins*100;
lfzmonate:lfzjahre*12;
barwert:(1-(1+monatszins)^(-1*lfzmonate))/monatszins;
rate:decimalplaces(kredit/barwert,2);
eqn:(1+x)^(lfzmonate)-decimalplaces(((kredit-gebuehr)/rate),2)*x*(1+x)^(lfzmonate)-1=0;roots:allroots(eqn);roots2:sublist(roots,lambda([t],imagpart(rhs(t))=0))effmonat:ev(x,roots[2]);xstart:decimalplaces(effmonat+.005,3);
fxstart:decimalplaces(ev(lhs(eqn),x=xstart),5);
fabl:diff(lhs(eqn),x,1);
fablxstart:decimalplaces(ev(fabl,x=xstart),5);
ta1:decimalplaces(xstart-(fxstart/fablxstart),5);
fehler:abs(decimalplaces(ev(lhs(eqn),x=ta1),5));
ta2:fehler;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </specificfeedback>
    <questionnote>
      <text>{@eqn@}{@xstart@}</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text />
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text />
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text />
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>1</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.8000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>getestet</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
    </tags>
  </question>

