<question type="stack">
    <name>
      <text>002) Ableitung</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Leiten Sie die Funktion \(f(x)={@f@}\)&lt;span class="latex"&gt;&lt;/span&gt; nach \(x\) ab.&lt;/p&gt;
&lt;p style="text-align: right;"&gt;(1 Punkt)&lt;/p&gt;
&lt;hr /&gt;
&lt;p style="-qt-block-indent: 0; text-indent: 0px; margin: 0px;"&gt;&lt;em&gt;Geben Sie die Antwort nach Möglichkeit ohne Verwendung von Leerzeichen in dem entsprechenden digitalen Antwortfeld ein.&lt;/em&gt;&lt;/p&gt;
&lt;p style="-qt-block-indent: 0; text-indent: 0px; margin: 0px;"&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;&lt;em&gt;Mit dem Haken neben dem Eingabefeld können Sie überprüfen, wie Ihr Ergebnis vom System gelesen wird. &lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/em&gt;&lt;/p&gt;
&lt;p style="-qt-block-indent: 0; text-indent: 0px; margin: 0px;"&gt;&lt;em&gt;Beispiele für Eingaben: &lt;/em&gt;\(3x \rightarrow 3x, \log(3x) \rightarrow \ln(3x), x \hat{} 3 \rightarrow x^3, x \hat{} (1/3) \rightarrow \sqrt[3]{x}, sqrt(3) \rightarrow \sqrt{3}, pi \rightarrow \pi \)&lt;br /&gt;&lt;br /&gt;&lt;/p&gt;
&lt;p style="-qt-block-indent: 0; text-indent: 0px; margin: 0px;"&gt;&lt;/p&gt;
&lt;p style="-qt-block-indent: 0; text-indent: 0px; margin: 0px;"&gt;&lt;strong&gt;Antwort:&lt;/strong&gt;&lt;/p&gt;
&lt;p&gt;\(f'(x)=\ \) [[input:Ableitung]] [[validation:Ableitung]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text />
    </generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>/* randomisierte Parameter */a: rand(4)+2;b: rand_with_prohib(-5,5,[0]);c: rand([3,6,9]);d: rand_with_step(3,9,2);/* Funktion */f: a*(x+b)^c+d;/* erste Ableitung der Funktion */M_Ableitung: diff(f,x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </specificfeedback>
    <questionnote>
      <text>\[\frac{d}{d{@x@}\left({@f@}\right)={@diff(f,x)@}\].</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct answer, well done.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;Your answer is partially correct.&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;Incorrect answer.&lt;/p&gt;</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>Ableitung</name>
      <type>algebraic</type>
      <tans>M_Ableitung</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>Ableitung</sans>
        <tans>M_Ableitung</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty>0.0000000</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty>0.0000000</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>ungetestet</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
      <tag><text>Faktorregel</text>
</tag>
    </tags>
  </question>

