<question type="stack">
    <name>
      <text>Geradengl Tangente</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;&lt;/p&gt;&lt;h4&gt;&lt;span style="font-size: 14px; font-weight: normal;"&gt;Sei \( f(x)= {@a@} x^{@b@} e^{{@n@}x}\).&lt;/span&gt;&lt;br&gt;&lt;/h4&gt;&lt;p&gt;Wie lautet die Geradengleichung der Tangente im Punkt \(x_0=1\)?&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;div class="editor-indent" style="margin-left: 30px;"&gt;&lt;p&gt;\(y=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]&lt;/p&gt;&lt;/div&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h4&gt;&lt;strong&gt;Lösung&lt;/strong&gt;&lt;/h4&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;Die Geradengleichung der Tangente \(f(x)\) bei \(x_0\) ist allgemein:&lt;/p&gt;&lt;div class="editor-indent" style="margin-left: 30px;"&gt;&lt;p&gt;\(y=f(x_0)+f^{\prime}(x_0)(x-x_0)\)&lt;/p&gt;&lt;/div&gt;&lt;p style="font-size: 17.5px;"&gt;&lt;span style="font-size: 17.5px;"&gt;&lt;br&gt;&lt;/span&gt;&lt;/p&gt;Hier:&lt;br&gt;&lt;p&gt;\(f(x)= {@a@} x^{@b@} e^{{@n@}x}\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;\(f^{\prime}(x)= ({@a*b@} x^{@b-1@} e^{{@n@}x})+({@a*n@} x^{@b@} e^{{@n@}x})\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;Setzte \(x_0=1\):&lt;/p&gt;&lt;p&gt;\(f(1)= {@a@} \cdot 1^{@b@} e^{{@n@}\cdot 1}={@a@}e^{{@n@}}\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;\(f^{\prime}(1)= ({@a*b@} \cdot 1^{@b-1@} e^{{@n@}\cdot 1}) + ({@a*n@} \cdot 1^{@b@} e^{{@n@}\cdot 1})={@a*b+a*n@}e^{@n@}\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;&lt;p&gt;Damit folgt aus der Formel:&lt;/p&gt;&lt;p&gt;\(y=f(1)+f^{\prime}(1)(x-1)={@a@}e^{{@n@}}+{@a*b+a*n@}e^{@n@}(x-1)={@a-a*b-a*n@}e^{{@n@}}+{@a*b+a*n@}xe^{@n@}\)&lt;br&gt;&lt;/p&gt;&lt;p&gt;&lt;br&gt;&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2020120600</text>
    </stackversion>
    <questionvariables>
      <text>a: 2+rand(8);
b:2 +rand(6);
n:2+rand(4);
fx0:a*e^n;
fstrichx0:a*b*e^n+ a*n*e^n;
y: fx0+ fstrichx0*(x-1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text />
    </specificfeedback>
    <questionnote>
      <text>{@a@}</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>y</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>y</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>getestet</text>
</tag>
      <tag><text>Lösung</text>
</tag>
      <tag><text>Tangente</text>
</tag>
      <tag><text>Übung</text>
</tag>
      <tag><text>Differentialrechnung</text>
</tag>
    </tags>
  </question>

